```python
# Initialize Otter
import otter
grader = otter.Notebook("discussion.ipynb")
```

# Discussion 05

### Due Saturday October 29th, 11:59:59PM

**Discussions will be due by the end of the day on Saturday** 

## Agenda: Midterm Review ✍️

* **Understand:** Take a look at Previous year midterm pattern and how this time's midterm would differ
* **Recap:** Revise/Recollect some lecture content related to the midterm
* **Practice:** Solve past year midterm problems for practice

```python
import pandas as pd
import seaborn as sns
import matplotlib
import matplotlib.pyplot as plt
import numpy as np
```

```python
from discussion import *
```

# Q1 - 0 Points

### The table below shows data collected for 330 penguins in Antarctica. Each row represents one penguin. 

### Three species of penguins appear in the data set: `Adelie`, `Gentoo`, and `Chinstrap`. This data set will be referred to throughout the exam as the DataFrame `df` .

```python
df = pd.read_csv('data/data1.csv')
df.dropna(inplace=True)
```

```python
df
```

# Q2 - 1 Point

## What kind of data is in the column named "Species"?
 - quantitative ?
 - nominal ?
 - ordinal ?

### Ans: 

# Q3 - 1 Point

## What kind of data is in the column named "Body Mass (g)"?
 - quantitative ?
 - nominal ?
 - ordinal ?

### Ans: 

# Q4 - 1 Point

## The code below produces a new table. How many rows will the table have? (See the dataset description above for useful information).


### `df.groupby('Species')['Culmen Length (mm)'].aggregate([np.mean, np.median])`

### Ans: 

```python
df.groupby('Species')['Culmen Length (mm)'].aggregate([np.mean, np.median])
```

# Q5 - 1 Point

## The following function accepts a series of numbers and returns a series of the same size in which every number has been standardized (also known as "z-scored"):

```python
def standardize(ser):
    return (ser - ser.mean()) / ser.std()
```


### Which of the following lines of code will return a series of standardized flipper lengths of length 330 (same as the number of rows in df ), where the standardization is done within each species?

### That is, the flipper sizes of Adelie penguins are standardized as a group, the flipper sizes of Gentoo penguins are standardized as another group, and so on.

### 1.  `df.groupby('Species')['Flipper Length (mm)'].transform(standardize)`
<!-- &nbsp;&nbsp; -->
### 2.  `df.groupby('Species')['Flipper Length (mm)'].aggregate(standardize)`
<!-- &nbsp;&nbsp; -->
### 3.  `df.groupby('Species')['Flipper Length (mm)'].standardize()`
<!-- &nbsp;&nbsp; -->
### 4.  `df.groupby('Species')['Flipper Length (mm)'].agg(standardize)`

### Ans: 

```python
df.groupby('Species')['Flipper Length (mm)'].transform(standardize)
```

```python
# df.groupby('Species')['Flipper Length (mm)'].aggregate(standardize)
```

```python
# df.groupby('Species')['Flipper Length (mm)'].standardize()
```

```python
# df.groupby('Species')['Flipper Length (mm)'].agg(standardize)
```

# Q6 - 1 Point

## Suppose incomplete data is collected on three new penguins, shown below:

```python
df_incomplete = pd.DataFrame({
    'Flipper Length (mm)':[191.0,193.0,197.0],
    'Culmen Depth (mm)':[20.3,17.7,19.5],
    'Species': ['Adelie','Adelie','Adelie']
})
df_incomplete
```

### Call this dataframe `df_incomplete`, and note that there are fewer columns than in `df` and that the columns which do appear are in a different order.

### Suppose `df_incomplete` is appended to the end of `df` with the code `pd.concat([df, df_incomplete])`. What will be the values in the last line of the resulting dataframe?

#### 1.     [197, 19.5, 'Adelie', nan, nan, nan, nan] with row label (index) 2 
#### 2.     [nan, 19.5, 197.0, nan, nan, nan, 'Adelie'] with row label (index) 2 
#### 3.     [0, 19.5, 197.0, 0, 0, 0, 'Adelie'] with row label (index) 2 
#### 4.     [197, 19.5, 'Adelie', 0, 0, 0, 0] with row label (index) 2 
#### 5.     [197, 19.5, 'Adelie', nan, nan, nan, nan] with row label (index) 332 
#### 6.     [nan, 19.5, 197.0, nan, nan, nan, 'Adelie'] with row label (index) 332 
#### 7.     [0, 19.5, 197.0, 0, 0, 0, 'Adelie'] with row label (index) 332 
#### 8.     [197, 19.5, 'Adelie', 0, 0, 0, 0] with row label (index) 332

### Ans: 

```python
df_concat = pd.concat([df,df_incomplete])
df_concat.tail()
```

# Q7 - 2 Points

### Suppose a penguin is considered "little" if its body mass is less than 3250 grams. Write a piece of code that computes the number of each species which are considered little. Your code should return a series. To receive full credit, your code should not modify the dataframe df .

### Ans: 

```python

```

### Note: your code will be graded manually, and it is not expected to be perfect. Be careful to not spend too much time trying to make your code perfect!

```python
df
```

# Q8 - 1 Point
### Suppose you have created a new table by_size which lists each penguin's size as either "small", "medium", or "large". The species is also included. The table looks like this:

<img src = "imgs/sizes.png"></img>

### Using the code `by_size.pivot_table(index='species', columns='size?', aggfunc='size')` , you've created the pivot table shown below:
<img src = 'imgs/pt.png'></img>

### There is a `NaN`in this pivot table. Why?

#### 1. some Gentoo penguins had sizes that were NaN
#### 2. some Gentoo penguins had sizes that were not strings
#### 3. there were too many small Gentoo penguins for the computer to represent with finite precision
#### 4. there were no small Gentoo penguins

### Ans: 

# Q9 - 1 Point 

### Suppose we wish to fill in the missing value in the pivot table above. What makes the most sense to fill it in with?

#### A. 0
#### B. 1
#### C. np.inf (infinity)
#### D. the mean of the "small" column
#### E. the mean of the "Gentoo" row
#### F. a number chosen at random from the observed values in the "small" column
#### G. a number at random from the uniform distribution on 0, 1

### Ans: 

# Q10 - 1 Point

### You have a second dataframe, `personalities`, that contains the disposition and intelligence of several species of penguin:
<img src = 'imgs/personalities.png'></img>

```python
personalities = pd.DataFrame({
    'Disposition':['happy','grumpy','mean'],
    'Intelligence':['smart','not smart','not smart'],
    'Species': ['Adelie','Chinstrap','Emperor']
})
personalities
```

### Note that personalities is missing some of the species that are in `df`, and it contains species that are not in `df`. You'd like to use the table to add columns to `df` for the disposition and intelligence of each of the penguins in the data set. 

### If a species in `df` is not listed in `personalities`, you're OK with having their disposition and Intelligence be `NaN`. 

### And if a species in `personalities` is not in `df`, it should not be added. 
### Therefore, the result of your merge should have exactly 330 rows -- the same as `df`.

### Which of the below will perform this?

#### 1. `pd.merge(df, personalities, how='outer')`
#### 2. `pd.merge(df, personalities, how='inner')`
#### 3. `pd.merge(df, personalities, how='left')`
#### 4. `pd.merge(df, personalities, how='right')`

### Ans: 

```python
pd.merge(df, personalities, how='outer')
```

```python
pd.merge(df, personalities, how='inner')
```

```python
pd.merge(df, personalities, how='left')
```

```python
pd.merge(df, personalities, how='right')
```

```python
df # Reminder that merge operations are not automatically inPlace
```

# Q11 Hypothesis Test - 7 Points

#### In the sample of penguins in the table `df` above, 40% are `Gentoo` penguins, 30% are `Adelie`, and 30% are `Chinstrap`.

#### It is known that, overall, 45% of penguins in Antarctica are `Gentoo`, 35% are `Adelie`, and 20% are `Chinstrap`. It therefore seems that the distribution of penguin species in your sample may differ from the population distribution.

#### You will test this with a hypothesis test. Your null hypothesis is that the sample in `df` was drawn at random from the population. The alternative hypothesis is that your sample was drawn with different probabilities than those shown above.

#### Write a piece of code to calculate the p-value of the observed data. You must choose a test statistic, calculate the observed value of the test statistic, compute simulated values of the test statistic, and calculate a p-value. For that reason, this question is worth 7 points, and partial credit will be awarded for each of the above.


#### Note: your code will be graded manually, and it is not expected to be perfect. Be careful to not spend too much time trying to make your code perfect!
#### Another Note: if you wish, you may write your code elsewhere and paste it into the box below.

### Ans: 

```python
penguins_type = pd.DataFrame([['Gentoo', 0.40, 0.45],
                    ['Adelie', 0.30, 0.35],
                    ['Chinstrap', 0.30, 0.20]],
                   columns=['Penguin Type', 'Sample', 'Overall']).set_index('Penguin Type')

penguins_type
```

```python
penguins_type.plot(kind='barh', title='Penguin Type Distribution of Sample vs Overall');
```

```python
observed_tvd = np.sum(np.abs(penguins_type['Sample'] - penguins_type['Overall'])) / 2
observed_tvd
```

```python
# Simulate one instance - Keep null hypothesis in mind
np.random.multinomial(1000, penguins_type['Overall']) / 1000
```

```python
# Simulate 5000 instances
N = 5000
np.random.multinomial(1000, penguins_type['Overall'], size=N) / 1000
```

```python
null_dists = np.random.multinomial(1000, penguins_type['Overall'], size=N) / 1000
null_tvds = np.sum(np.abs(null_dists - penguins_type['Overall'].to_numpy()), axis=1) / 2
null_tvds
```

```python
pd.Series(null_tvds).plot(kind='hist', 
                     density=True,
                     ec='w',
                     title='Simulated Null TVDs & Observed TVD');
plt.axvline(x=observed_tvd, color='red', linewidth=2);
```

**Question 1**

Now that we've finalized our test statistic and since we've also simulated the null hypothesis under the test statistic, define a function `prev_midterm_hypothesis_test` that takes in the observations `penguins_type` and does the following

1. Calculate the `observed_tvd`
2. Calculate the `p_value` after simulating the null statistic on 5000 samples

The function should return a tuple with in the following format `(observed_tvd,p_val)`

```python
# don't change this cell -- it is needed for the tests to work
penguins_type = pd.DataFrame([['Gentoo', 0.40, 0.45],
                    ['Adelie', 0.30, 0.35],
                    ['Chinstrap', 0.30, 0.20]],
                   columns=['Penguin Type', 'Sample', 'Overall']).set_index('Penguin Type')
q1_out = prev_midterm_hypothesis_test(penguins_type)
```

```python
grader.check("q1")
```

# Q12 - 1 Point

### Suppose you look at the distribution of culmen (beak) depth for each species and find the following:

<img src = 'imgs/culmen.png'></img>

### It looks as though Chinstrap penguins have larger culmen depth than Gentoo penguins.

### Which of the following test statistics could be used to test this alternative hypothesis against the null hypothesis that the two distributions are the same? Check all that apply.

### 1. TVD
### 2. Kolmogorov-Smirnov
### 3. signed difference in mean culmen depth
### 4. unsigned difference in mean culmen depth

### Ans: 

# Q13 - 1 Point

### Suppose you now look at the distribution of culmen depth for each species once more and again find the following:

<img src = 'imgs/culmen1.png'></img>

### What test statistic is the best choice for testing whether the empirical distributions of culmen depth in Chinstrap and Adelie penguins came from different underlying distributions?

#### 1. the Kolmogorov-Smirnov Statistic
#### 2. the absolute difference in means
#### 3. the TVD
#### 4. the signed difference in means

### Ans: 

# Q14 - 1 Point
### Suppose some of the entries in the "Body Mass (g)" column are missing.


### Body mass measurements were taken any time a penguin walked across a scale that was placed outside. In your experience, the scale sometimes fails to make a measurement if the item placed on top is too small.

### What is the most likely type of missingness in this case?

#### 1. Missing Completely at Random
#### 2. Missing at Random
#### 3. Not Missing at Random
#### 4. Missing by Design

### Ans: 

# Q15 - 1 Point

### Suppose some of the entries in the "Delta 15 N (o/oo)" and "Delta 13 C (o/oo)" columns are missing.


### The measurements above were taken by four different teams located at different places in Antarctica, and one of the teams had a faulty measurement device that prevented them from measuring "Delta 15 N" and "Delta 13 N". The penguins in this team's area were mostly Gentoo penguins.

### What is the most likely type of missingness in this case?

#### 1. Missing Completely at Random
#### 2. Missing at Random
#### 3. Not Missing at Random
#### 4. Missing by Design

### Ans: 

# Q16 - 1 Point

### Suppose some of the entries in the "Flipper Length (mm)" column are missing. It is very hard to measure the flipper length of a wild penguin because penguins have a tendency to run away whenever you approach them with a ruler. 

### For this reason, the researchers agree that they will only measure the flipper length of half of the penguins they see. To implement this, they flip a coin when they see a penguin -- if the coin comes up heads, they measure the penguin's flipper length, and if it is tails they leave the flipper length blank.

### What is the most likely type of missingness in this case?


#### 1. Missing Completely at Random
#### 2. Missing at Random
#### 3. Not Missing at Random
#### 4. Missing by Design

### Ans: 

# Q17 - 1 Point

### Suppose some of the body masses are missing. You'd like to test if some of the body masses are missing at random (MAR) dependent on the species. 

### You will determine this with a permutation test. Which one of the below is a valid test statistic for this test?


#### 1. The TVD (total variation distance)
#### 2. The K-S statistic (Kolmogorov-Smirnov)
#### 3. The absolute (unsigned) difference in means
#### 4. The signed difference in means

# Q18 - 1 Point


### Suppose the "Flipper Length (mm)" column has missing values. The distributions of species when the flipper length is missing (null) and when it is present (not null) are shown below:

<img src = 'imgs/missingness_distribution.png'></img>

### Futhermore, the average flipper length of each species is as follows:

```python
df.groupby('Species')['Flipper Length (mm)'].mean()
```

### Suppose the overall mean flipper length is computed (using only the observed values).
### Which of the following is true?

### A. the mean will be biased high
### B. the mean will be biased low
### C. the mean will be unbiased

### Ans: 

# Q19 - 2 Points

### Write a piece of code that imputes missing flipper lengths with the median flipper length of the species that the penguin belongs to.

**Question 2**

Write a helper function, `midterm_review_imputation` that accepts `penguins_with_nulls` dataframe as an argument and returns a `pd.DataFrame` object where the cells in which `Flipper Length (mm)` is missing are **imputed with the median flipper length of the species** that the penguin belongs to. 



```python
# Do Not change this cell, it is needed for the autograder code to run correctly
penguins_with_nulls = pd.read_csv('data/data1.csv')
q2_out = midterm_review_imputation(penguins_with_nulls)
```

```python
grader.check("q2")
```

## Congratulations! You're done!

* Submit your `.py` file to Gradescope. Note that you only need to submit the `.py` file; this notebook should not be uploaded. Make sure that all of your work is in the `.py` file and not here by running the doctests: `python -m doctest discussion.py`.

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for q1 (points: None)
```
>>> isinstance(q1_out, tuple)
True
```
failure_message: doctest
```
>>> 0<=q1_out[0]<=1
True
```
failure_message: Error in observed TVD
```
>>> np.isclose(q1_out[0],0.09999999999999998,atol=0.1)
True
```
failure_message: Error in observed TVD
```
>>> np.isclose(q1_out[0],0.09999999999999998,atol=0.01)
True
```
failure_message: Error in observed TVD
```
>>> isinstance(q1_out[1], float)
True
```
failure_message: Error in observed P-VALUE


Tests for q2 (points: None)
```
>>> isinstance(q2_out, pd.DataFrame)
True
```
failure_message: doctest
```
>>> q2_out['Flipper Length (mm)'].isnull().sum() == 0
True
```
failure_message: doctest
```
>>> q2_out.loc[3,'Flipper Length (mm)'] == 190.0
True
```
failure_message: Null Flipper Length not correctly imputed at some indices
```
>>> q2_out.isnull().sum().sum() == 33
True
```
failure_message: Null Flipper Length not correctly imputed
```
>>> q2_out.loc[339,'Flipper Length (mm)'] == 216.0
True
```
failure_message: Null Flipper Length not correctly imputed at some indices
